---
  title: "Introduction to Machine Learning"
  description: "This final chapter will give you a taste of the all-powerful Machine Learning. You will learn how to preprocess your data and apply a Decision Tree algorithm to predict the home planet of a given customer."
---

## Martian Casino

```yaml
type: TabExercise 
xp: 100 
key: c864253af5   
```


A new casino has opened on mars, and they need your help! 

In one of the games, a customer pays 1 murzel, the local martian currency, to roll 2 dice. If they can roll a 10 or higher, they win 5 murzels. However, the low gravity on Mars means that the dice keep bouncing away! 

The casino would like to trial a computer version, in which 2 random numbers between 1 and 6 are generated instead of using real dice. A customer called Zworby has volunteered to trial the game, and would like to play 100 times.


`@pre_exercise_code`

```{python}
#none
```

`@sample_code`

```{python}
#none
```

***



```yaml
type: NormalExercise 
xp: 100 
key: b191179b95   
```





`@instructions`


`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


