---
  title: "Introduction to Machine Learning"
  description: "This final chapter will give you a taste of the all-powerful Machine Learning. You will learn how to preprocess your data and apply a Decision Tree algorithm to predict the home planet of a given customer."
---

## Martian Casino

```yaml
type: TabExercise 
xp: 100 
key: c864253af5   
```


A new casino has opened on mars, and they need your help! 

In one of the games, a customer pays 1 murzel, the local martian currency, to roll 2 dice. If they can roll a 10 or higher, they win 5 murzels. However, the low gravity on Mars means that the dice keep bouncing away! 

The casino would like to trial a computer version, in which 2 random numbers between 1 and 6 are generated instead of using real dice. A customer called Zworby has volunteered to trial the game, and would like to play 100  times.


`@pre_exercise_code`

```{python}
#none
```

`@sample_code`

```{python}
#none
```

***



```yaml
type: NormalExercise 
xp: 100 
key: b191179b95   
```





`@instructions`
Generate 100 sets of 2 dice rolls, adding their total to a list.

In Python, often the best way to learn how to do something, is to google it! There's most likely many people who will have attempted the same thing.

Stack overflow is a great site, where over 16 million questions have been asked, and members of the community help each other out. The link below is a question on how to make a random number generator. Change their example to make your number generator to replace a dice. 

[https://stackoverflow.com/questions/3996904/generate-random-integers-between-0-and-9](url)

`@hint`


`@sample_code`

```{python}

```


`@solution`

```{python}

```


`@sct`

```{python}

```


