---
  title: "Introduction to Machine Learning"
  description: "This final chapter will give you a taste of the all-powerful Machine Learning. You will learn how to preprocess your data and apply a Decision Tree algorithm to predict the home planet of a given customer."
---

## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: cb1eac1984   
```


a


`@instructions`
a

`@hint`


`@pre_exercise_code`

```{python}

```


`@sample_code`

```{python}
print('a')
```

`@solution`

```{python}
print('a')
```

`@sct`

```{python}
Ex().has_equal_ast()
success_msg("Well done!")
```
